# Demo Code Challenge 

This demo is designed to give you the experience of a Code Challenge before your official Code Challenge.
 
*Read the instructions carefully.* Your code will need to meet detailed specifications to pass automated tests.

## Code Tests

We have provided some code tests for you to run to check that your work meets the item specifications. Passing these tests does not necessarily mean that you have gotten the item correct - there are additional hidden tests. However, if any of the tests do not pass, this tells you that your code is incorrect and needs changes to meet the specification. To determine what the issue is, read the comments in the code test cells, the error message you receive, and the item instructions.

## 1) Read `titanic.csv` into a pandas DataFrame named `df`

Use pandas to create a new DataFrame, called `df`, containing the data from the dataset in the file `titanic.csv` in the folder containing this notebook. 

Hint: Use the string `'./titanic.csv'` as the file reference.

```python
! ls titanic.csv
```

Output:
```
titanic.csv
```

```python
# Run this cell without changes

import pandas as pd
import numpy as np
from numbers import Number
import warnings
warnings.filterwarnings('ignore')
```

```python
# CodeGrade step1
# Replace None with appropriate code

df = pd.read_csv('./titanic.csv')
df.head()
```

Output:
```
PassengerId  Survived  Pclass  \
0            1         0       3   
1            2         1       1   
2            3         1       3   
3            4         1       1   
4            5         0       3   

                                                Name     Sex   Age  SibSp  \
0                            Braund, Mr. Owen Harris    male  22.0      1   
1  Cumings, Mrs. John Bradley (Florence Briggs Th...  female  38.0      1   
2                             Heikkinen, Miss. Laina  female  26.0      0   
3       Futrelle, Mrs. Jacques Heath (Lily May Peel)  female  35.0      1   
4                           Allen, Mr. William Henry    male  35.0      0   

   Parch            Ticket     Fare Cabin Embarked  
0      0         A/5 21171   7.2500   NaN        S  
1      0          PC 17599  71.2833   C85        C  
2      0  STON/O2. 3101282   7.9250   NaN        S  
3      0            113803  53.1000  C123        S  
4      0            373450   8.0500   NaN        S
```

<!-- grade_id: cell-f7427f34ad14af7a, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a DataFrame named df

assert type(df) == pd.DataFrame
```

### 2) Create a variable `dogwood_index` containing the index for the `'dogwood'` element from the list `tree_type_names`

Below is the list `tree_type_names` that you will need - do not change this list.


```python
# Run this cell without changes

tree_type_names = ['linden', 'spruce', 'dogwood', 'hickory', 'willow']
```

```python
# CodeGrade step2
# Replace None with appropriate code

dogwood_index = tree_type_names.index('dogwood')
dogwood_index
```

Output:
```
2
```

<!-- grade_id: cell-00926387ac95d868, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named dogwood_index

assert isinstance(dogwood_index, Number)
```

## 3)  Briefly explain the purpose of code challenges in the Flatiron School Data Science program.

Enter a one sentence response below.

```python
# Your answer here
"The code challenge lets the instructor know how everyone is doing in the course so far"

```

Output:
```
'The code challenge lets the instructor know how everyone is doing in the course so far'
```

```python

```